# CSCI 3155 Recitation 6
October 5, 2018

## Closures

https://docs.google.com/presentation/d/1H5coQ04RiNxKWTAb6ljWDiCJKqk_2w52SaosYGVqZ38/edit?usp=sharing

## Mutual Recursion

Recursion is defined as a function calling itself in the function body:
```
def f_rec(): Unit = f_rec()
```

Mutual recursion is when a function calls a helper function that then makes the call back to it:
```
def f_mutual_rec_1(): Unit = f_mutual_rec_2()
def f_mutual_rec_2(): Unit = f_mutual_rec_1()
```

Below is an example of a useful function that makes use of mutual recursion.

```scala
def isEvenLen(l: List[Int]): Boolean = l match {
    case _ :: rest => isOddLen(rest)
    case Nil => true
}

def isOddLen(l: List[Int]): Boolean = l match {
    case _ :: rest => isEvenLen(rest)
    case Nil => false
}
```

Output:
```
defined [32mfunction[39m [36misEvenLen[39m
defined [32mfunction[39m [36misOddLen[39m
```

```scala
val l_0 = Nil
assert(isEvenLen(l_0))
assert(!isOddLen(l_0))

val l_1 = 2 :: l_0
assert(isOddLen(l_1))
assert(!isEvenLen(l_1))

val l_2 = 1 :: l_1
assert(isEvenLen(l_2))
assert(!isOddLen(l_2))
```

Output:
```
[36ml_0[39m: [32mNil[39m.type = [33mList[39m()
[36ml_1[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m2[39m)
[36ml_2[39m: [32mList[39m[[32mInt[39m] = [33mList[39m([32m1[39m, [32m2[39m)
```

Mutually recursive functions are often useful when the input has a mutually recursive structure as well. For example, we have seen binary trees defined inductively:
```
sealed trait BTree
case class EmptyTree() extends BTree
case class TreeNode(l: BTree, d: Int, r: BTree) extends BTree
```

As well as lists:
```
sealed trait List
case class EmptyList() extends List
case class ListNode(d: Int, t: List)
```

We could combine these two inductive structures, and create an n-ary trees _(shamelessly stolen fom Wikipedia)_:

```scala
sealed trait NaryTree
case object Empty extends NaryTree
// Rely on already defined inductive Lists
case class Node(d: Int, childen: List[NaryTree]) extends NaryTree
```

Output:
```
defined [32mtrait[39m [36mNaryTree[39m
defined [32mobject[39m [36mEmpty[39m
defined [32mclass[39m [36mNode[39m
```

```scala
def countNodesTree(tree: NaryTree): Int = tree match {
    case Empty => 0
    case Node(d, children) => 1 + countNodesChildren(children)
}

def countNodesChildren(trees: List[NaryTree]): Int = trees match {
    case Nil => 0
    case tree :: rest => countNodesTree(tree) + countNodesChildren(rest)
}
```

Output:
```
defined [32mfunction[39m [36mcountNodesTree[39m
defined [32mfunction[39m [36mcountNodesChildren[39m
```

```scala
val tree = Node(0,
                Node(1, 
                     Node(5, Nil) :: Nil ) ::
                Node(2, Nil) :: 
                Node(3, Nil) :: 
                Node(4, Nil) :: Nil
)

assert(countNodesTree(tree) == 6)
```

Output:
```
[36mtree[39m: [32mNode[39m = Node(0,List(Node(1,List(Node(5,List()))), Node(2,List()), Node(3,List()), Node(4,List())))
```

### Exercise: Check heapiness

Implement a function that checks whether a givenn NaryTree is a min-heap (A tree where every node has a smaller value than any children). 

<!-- grade_id: cell-4ba87e6d71c83a74, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
def isMinHeapTree(tree: NaryTree, min: Int = Int.MinValue): Boolean = tree match {
    case Empty => true
    case Node(d, children) if(d > min) => isMinHeapChildren(children, d)
}

def isMinHeapChildren(trees: List[NaryTree], min: Int = Int.MinValue): Boolean = trees match {
    case Nil => true
    case Node(d, children) :: rest if(d > min) => {
        val tree = Node(d, children)
        isMinHeapTree(tree, min) && isMinHeapChildren(rest, min)
    }
}


// def isMinHeapTree(tree: NaryTree, min: Int = Int.MinValue): Boolean = tree match {
//     case Empty => true
//     case Node(d, children) if(d > min) => isMinHeapChildren(children, d)
// }

// def isMinHeapChildren(trees: List[NaryTree], min: Int = Int.MinValue): Boolean = trees match {
//     case Nil => true
//     case tree :: rest => {
//         isMinHeapTree(tree, min) && isMinHeapChildren(rest, min)
//     }
// }
```

Output:
```
defined [32mfunction[39m [36misMinHeapTree[39m
defined [32mfunction[39m [36misMinHeapChildren[39m
```

<!-- grade_id: cell-1cfa2c1ebbc45899, points: 1, grade: true, solution: false -->
```scala
val tree_minHeap = Node(0,
                        Node(1, 
                             Node(5, Nil) :: Nil ) ::
                        Node(2, Nil) :: 
                        Node(3, Nil) :: 
                        Node(4, Nil) :: Nil)
assert(isMinHeapTree(tree_minHeap), "1")

val tree_NOT_minHeap = Node(5, tree_minHeap :: Nil)
assert(!isMinHeapTree(tree_NOT_minHeap), "2")
```

Output:
```
$sess.cmd6Wrapper.<init>(cmd6.sc:131)
  $sess.cmd6$.<init>(cmd6.sc:82)
  $sess.cmd6$.<clinit>(cmd6.sc:-1)
```

```scala

```