# Assignment 1: Scala Programming

This assignment asks you to write scala programs. Fill in the programs below. Use the test cases provided to test them. You are also strongly encouraged to contribute your own test cases.

## Problem 1 (5 points)

Write a scala function titled `isPalindrome (s)` that tests whether a given input string `s` is a palindrome.

The following string operations will be helpful.

1. `s.size` gives you the size of the string.
2. `s(i)` gives you the `ith` character for `i = 0, ... , s.size-1`
3. For loops in scala are described here: https://www.tutorialspoint.com/scala/scala_for_loop.htm
   Pay careful attention between the to and until forms of the for loop.

<!-- grade_id: a1, points: 0, grade: false, solution: true -->
```scala
// Julia Sanford

def isPalindrome(s: String): Boolean = {
    // YOUR CODE HERE
    var len = s.size - 1;
    var start = 0;
    while (start < len){
        if (s(start) != s(len)){
            return false;
        }
        start += 1;
        len -= 1;
    }
    return true;
}
```

Output:
```
defined [32mfunction[39m [36misPalindrome[39m
```

<!-- grade_id: t1, points: 5, grade: true, solution: false -->
```scala
// TEST CASES

assert(isPalindrome("hellolleh") == true, "1")
assert(isPalindrome("bababablacksheep") == false, "2")
assert(isPalindrome("") == true, "3")
assert(isPalindrome("a") == true, "4")
assert(isPalindrome("xxxxxxxx") == true, "5")
assert(isPalindrome("ablewasiereisawelba") == true, "6")
```

## Problem 2 (5 points)

Implement a `isPalindromeRec` function that checks if a given string is a palindrome, but without using mutable `var` or without using `for`/`while` loops.

**Hint:** You should consult Scala documentation on strings to figure out:
1. How do we get first and last characters of a string.
2. How do we get substring between first and last character.
3. How to find out if a string is empty or a single character.
https://www.tutorialspoint.com/scala/scala_strings.htm

<!-- grade_id: a2, points: 0, grade: false, solution: true -->
```scala
def isPalindromeRec(s: String): Boolean = {
    // YOUR CODE HERE
    val l = s.size
    if (l == 1 || l == 0 || (l == 2 && s(0) == s(1))){
        return true
    }
    val front = s(0)
    val end = s(l-1)
    if (front == end && l > 2){
        return isPalindromeRec(s.substring(1, l-1))
    }
    else {
        return false
    }
}
```

Output:
```
defined [32mfunction[39m [36misPalindromeRec[39m
```

<!-- grade_id: t2, points: 5, grade: true, solution: false -->
```scala
// TEST CASES

assert(isPalindromeRec("hellolleh") == true, "1")
assert(isPalindromeRec("bababablacksheep") == false, "2")
assert(isPalindromeRec("") == true, "3")
assert(isPalindromeRec("a") == true, "4")
assert(isPalindromeRec("xxxxxxxx") == true, "5")
assert(isPalindromeRec("ablewasiereisawelba") == true, "6")
```

## Problem 3 (5 points)

Write a scala function XOR to compute the XOR of two boolean inputs.

<!-- grade_id: a3, points: 0, grade: false, solution: true -->
```scala
def computeXOR(x1: Boolean, x2: Boolean): Boolean = {
    // YOUR CODE HERE
    if (x1 == x2){
        return false;
    }
    return true;
}
```

Output:
```
defined [32mfunction[39m [36mcomputeXOR[39m
```

<!-- grade_id: t3, points: 5, grade: true, solution: false -->
```scala
// TEST CASES

assert(!computeXOR(true, true), "1")
assert(computeXOR(true, false), "2")
assert(computeXOR(false, true), "3")
assert(!computeXOR(false, false), "4")
```

## Problem 4 (15 points)

We would like you to implement the methods in the class Rational that is supposed to represent fractions $\frac{n}{d}$, where the numerator and denominator are integers $n, d$. Note that it is important to ensure that $d \not= 0$.

Implement the missing methods below.

You should read this blog post about how to convert between numeric types.

https://alvinalexander.com/scala/how-to-convert-between-numeric-types-in-scala-int-long-float-double


<!-- grade_id: a4, points: 0, grade: false, solution: true -->
```scala
class Rational(val n: Int, val d: Int) {
    // Implement a pretty printer
    override def toString(): String = {
        s"$n/$d"
    }
    
    def isValid(): Boolean = {
        // check if this is a valid fraction.
        // YOUR CODE HERE
        if (d == 0){
            return false
        }
        true
    }
    def toDouble(): Double = {
        // Convert n/ d to double
        // use toDouble method to convert an integer to double
        // YOUR CODE HERE
        val num = n.toDouble
        val den = d.toDouble
        return num/den
    }
    
    def isInteger(): Boolean = {
        // Is this fraction an integer.
        // A fraction is an integer if n is divisible by d
        // YOUR CODE HERE
        if (n % d == 0){
            return true
        }
        false
    }
    
    def floor(): Double = {
        // return the floor of n/d
        // math.floor(f) takes the floor of a double precision number in scala
        // YOUR CODE HERE
        return math.floor(toDouble())
    }
    
    
    def rationalEquals(r: Rational): Boolean = {
        // Check if this equals r
        // YOUR CODE HERE
        if (toDouble() == r.toDouble()){
            return true
        }
        false
    }
    
    
    
    def + (r: Rational) : Rational = {
        // Add two rationals and return a new Rational
        // YOUR CODE HERE
        val n1 = n * r.d;
        val rn1 = r.n * d;
        val R: Rational = new Rational(n1 + rn1, r.d * d);
        return R
    }
    
    def - (r: Rational): Rational = {
        // compute this - r and return a new Rational
        // YOUR CODE HERE
        val n1 = n * r.d;
        val rn1 = r.n * d;
        val R: Rational = new Rational(n1 - rn1, r.d * d);
        return R
    }
    
    def * (r: Rational): Rational = {
        // multiply two rationals and return a new Rational
        // YOUR CODE HERE
        val R: Rational = new Rational(n * r.n, d * r.d)
        return R
    }
    
    def / (r: Rational): Rational = {
        // compute this/r and return a new Rational
        // YOUR CODE HERE
        val R: Rational = new Rational(n * r.d, d * r.n)
        return R
    }
    
    // We can use your rationalEquals method to overide the == operator as follows
    // Ignore this method for now -- it will be clearer soon
    override def equals(r: Any): Boolean = r match {
        case rat: Rational => this.rationalEquals(rat) 
        case _ => false
    }
    // END Ignore
    
}
```

Output:
```
defined [32mclass[39m [36mRational[39m
```

<!-- grade_id: t4_isValid, points: 1, grade: true, solution: false -->
```scala
// TEST CODE
assert(new Rational(10, 0).isValid() == false, "1 - isValid")
assert(new Rational(10, 1).isValid() == true, "2 - isValid")
```

<!-- grade_id: t4_toDouble, points: 1, grade: true, solution: false -->
```scala
// TEST CODE
assert(new Rational(10, 3).toDouble() == 10.0/3.0, "1 - toDouble")
assert(new Rational(19, 5).toDouble() == 19.0/5.0, "2 - toDouble")
```

<!-- grade_id: t4_isInteger, points: 1, grade: true, solution: false -->
```scala
// TEST CODE
assert(new Rational(15, 3).isInteger() == true, "1 - isInteger")
assert(new Rational(15, 8).isInteger() == false, "2 - isInteger")
assert(new Rational(225, 15).isInteger() == true, "3 - isInteger")
assert(new Rational(181, 2).isInteger() == false, "4 - isInteger")
```

<!-- grade_id: t4_floor, points: 2, grade: true, solution: false -->
```scala
// TEST CODE
assert(new Rational(15, 4).floor() == 3.0, "1 - floor")
assert(new Rational(100, 45).floor() == 2.0, "2 - floor")
assert(new Rational(-25, 3).floor() == -9.0, "3 - floor")
assert(new Rational(-40, 7).floor() == -6.0, "4 - floor")
```

<!-- grade_id: t4_eq, points: 2, grade: true, solution: false -->
```scala
// TEST CODE
assert((new Rational(20, 5)) == (new Rational(16 , 4)), "1 - rationalEquals")
assert((new Rational(24, 5)) == (new Rational(120, 25)), "2 - rationalEquals")
assert((new Rational(19, 5)) != (new Rational(5, 19)), "3 - rationalEquals")
assert((new Rational(24, 1)) != (new Rational(120, 15)), "4 - rationalEquals")
```

<!-- grade_id: t4_plus, points: 2, grade: true, solution: false -->
```scala
// TEST CODE
val r1 = new Rational(10, 3)
val r2 = new Rational(14, 3)
val r3 = r1 + r2
assert(r3.n == r3.d * 8, "1) addition of '10/3 + 14/3'")

val r4 = new Rational(9, 2)
val r5 = new Rational(18, 2)
val r6 = r4 + r5
assert(r6.n * 2 == r6.d * 27, "2) addition of '9/2 + 18/2'")

val r7 = new Rational(15, 7)
val r8 = new Rational(19, 6)
val r9 = r7 + r8
assert(r9.n * 42 == r9.d * 223, "3) addition of '15/7 + 19/6'")
```

Output:
```
[36mr1[39m: [32mRational[39m = 10/3
[36mr2[39m: [32mRational[39m = 14/3
[36mr3[39m: [32mRational[39m = 72/9
[36mr4[39m: [32mRational[39m = 9/2
[36mr5[39m: [32mRational[39m = 18/2
[36mr6[39m: [32mRational[39m = 54/4
[36mr7[39m: [32mRational[39m = 15/7
[36mr8[39m: [32mRational[39m = 19/6
[36mr9[39m: [32mRational[39m = 223/42
```

<!-- grade_id: t4_minus, points: 2, grade: true, solution: false -->
```scala
// TEST CODE
val r1 = new Rational(10, 3)
val r2 = new Rational(13, 3)
val r3 = r1 - r2
assert(r3.n == -r3.d, "1) subtraction of '10/3 - 13/3'")

val r4 = new Rational(18, 2)
val r5 = new Rational(9, 2)
val r6 = r4 - r5
assert(r6.n * 2 == r6.d * 9, "2) subtraction of '18/2 - 9/2'")

val r7 = new Rational(15, 7)
val r8 = new Rational(19, 6)
val r9 = r7 - r8
assert(r9.n * 42 == r9.d * -43, "3) subtraction of '15/7 - 19/6'")
```

Output:
```
[36mr1[39m: [32mRational[39m = 10/3
[36mr2[39m: [32mRational[39m = 13/3
[36mr3[39m: [32mRational[39m = -9/9
[36mr4[39m: [32mRational[39m = 18/2
[36mr5[39m: [32mRational[39m = 9/2
[36mr6[39m: [32mRational[39m = 18/4
[36mr7[39m: [32mRational[39m = 15/7
[36mr8[39m: [32mRational[39m = 19/6
[36mr9[39m: [32mRational[39m = -43/42
```

<!-- grade_id: t4_times, points: 2, grade: true, solution: false -->
```scala
// TEST CODE
val r1 = new Rational(10, 3)
val r2 = new Rational(13, 3)
val r3 = r1 * r2
assert(r3.n * 9 == r3.d * 130, "1) multiplication of '10/3 * 13/3'")

val r4 = new Rational(-18, 2)
val r5 = new Rational(9, 2)
val r6 = r4 * r5
assert(r6.n * -2 == r6.d * 81, "2) multiplication of '-18/2 * 9/2'")

val r7 = new Rational(15, 7)
val r8 = new Rational(19, 6)
val r9 = r7 * r8
assert(r9.n * 14 == r9.d * 95, "3) multiplication of '15/7 * 19/6'")
```

Output:
```
[36mr1[39m: [32mRational[39m = 10/3
[36mr2[39m: [32mRational[39m = 13/3
[36mr3[39m: [32mRational[39m = 130/9
[36mr4[39m: [32mRational[39m = -18/2
[36mr5[39m: [32mRational[39m = 9/2
[36mr6[39m: [32mRational[39m = -162/4
[36mr7[39m: [32mRational[39m = 15/7
[36mr8[39m: [32mRational[39m = 19/6
[36mr9[39m: [32mRational[39m = 285/42
```

<!-- grade_id: t4_div, points: 2, grade: true, solution: false -->
```scala
// TEST CODE
val r1 = new Rational(10, 3)
val r2 = new Rational(13, 3)
val r3 = r1 / r2
assert(r3.n * 13 == r3.d * 10, "1) division of '(10/3) / (13/3)'")

val r4 = new Rational(-18, 2)
val r5 = new Rational(9, 3)
val r6 = r4 / r5
assert(r6.n == r6.d * -3, "2) division of '(-18/2) / (9/3)'")

val r7 = new Rational(15, -7)
val r8 = new Rational(19, -6)
val r9 = r7 / r8
assert(r9.n * 133 == r9.d * 90, "3) division of '(15/-7) / (19/-6)'")
```

Output:
```
[36mr1[39m: [32mRational[39m = 10/3
[36mr2[39m: [32mRational[39m = 13/3
[36mr3[39m: [32mRational[39m = 30/39
[36mr4[39m: [32mRational[39m = -18/2
[36mr5[39m: [32mRational[39m = 9/3
[36mr6[39m: [32mRational[39m = -54/18
[36mr7[39m: [32mRational[39m = 15/-7
[36mr8[39m: [32mRational[39m = 19/-6
[36mr9[39m: [32mRational[39m = -90/-133
```

## Problem 5

Newton invented the Newton-Raphson method for solving an equation. We are going to ask you to write some code to solve equations.

To solve an equation of the form

$$ x^2 - 3x + 2 == 0$$

we start from an initial guess at the solution: say $$x_0 = 4.5$$

Each time we have the $i^th$ guess $x_i$, we update it as

$$ x_{i+1} = x_i - \frac{f(x_i)}{f'(x_i)} $$

For our equation, $f(x) = x^2 - 3x +2$ and $f'(x) = 2 x - 3$.

Thus, our update equation is 
$$ x_{i+1} = x_i - \frac{x_i^2 - 3 x_i + 2}{2 x_i - 3}$$.

We stop whenever $|f(x_i)| \leq 10^{-8}$: i.e, we are very close to a root of the function.

Gory details are here:
http://www.math.ubc.ca/~anstee/math104/newtonmethod.pdf

__(A, 5 points)__ Write a scala function `calculateRoot(x0: Double)` that takes in the initial guess $x_0$ as input and calculates the root for the polynomial $x^2 - 3 x + 2$. Your code should use a While loop.

Caution: To calculate $x^2$, you should use `math.pow(x,2)` in Scala.

<!-- grade_id: a5a, points: 0, grade: false, solution: true -->
```scala
def calculateRoot(x0: Double): Double = {
    def f(x: Double) = { math.pow(x,2) - 3 * x + 2.0 }
    def df(x: Double) = { 2* x - 3}
    // YOUR CODE HERE
    var xi = x0
    while (math.abs(f(xi)) > math.pow(10, -8)){
        xi = xi - (f(xi) / df(xi))
    }
    xi;
}
```

Output:
```
defined [32mfunction[39m [36mcalculateRoot[39m
```

<!-- grade_id: t5a, points: 5, grade: true, solution: false -->
```scala
// TEST CODE

def isRootCorrect(x: Double): Boolean = {
    val y = calculateRoot(x)
    (math.abs(y - 2.0) <= 1e-06) || (math.abs(y - 1.0) <= 1e-06)
}

assert(isRootCorrect(3.5), "1")
assert(isRootCorrect(4.6), "2")
assert(isRootCorrect(-3), "3")
assert(isRootCorrect(-3.5), "4")
```

Output:
```
defined [32mfunction[39m [36misRootCorrect[39m
```

__(B, 5 points)__ Implement the function `calculateRootRec` that uses a recursion instead of a while loop to compute the root. You should __not__ use a mutable `var` in your code. 

<!-- grade_id: a5b, points: 0, grade: false, solution: true -->
```scala
def calculateRootRec(x: Double): Double = {
    // YOUR CODE HERE
    if (math.abs(math.pow(x,2) - 3 * x + 2.0) > math.pow(10, -8)){
        return calculateRootRec(x - ((math.pow(x,2) - 3 * x + 2.0) / (2* x - 3)))
    }
    x
}
```

Output:
```
defined [32mfunction[39m [36mcalculateRootRec[39m
```

<!-- grade_id: t5b, points: 5, grade: true, solution: false -->
```scala
// TEST CODE
def isRootCorrectRec(x: Double): Boolean = {
    val y = calculateRootRec(x)
    (math.abs(y - 2.0) <= 1e-06) || (math.abs(y - 1.0) <= 1e-06)
}

assert(isRootCorrectRec(3.5), "1")
assert(isRootCorrectRec(4.6), "2")
assert(isRootCorrectRec(-3), "3")
assert(isRootCorrectRec(-3.5), "4")
assert(isRootCorrectRec(1.3), "5")
```

Output:
```
defined [32mfunction[39m [36misRootCorrectRec[39m
```

__(C, 5 points)__ Write a generalized Newton-Raphson function `solveEquationNewtonRaphson` to solve any 
equation $f(x) = 0$ wherein the functions $f$ and $f'(x)$ are given as inputs. Do not use while loops or mutable `var` in your code.



<!-- grade_id: a5c, points: 0, grade: false, solution: true -->
```scala
def solveEquationNewtonRaphson(x: Double, f: Double => Double, df: Double => Double): Double = {
    // YOUR CODE HERE
    if (math.abs(f(x)) > math.pow(10, -8)){
        return solveEquationNewtonRaphson(x - (f(x) / df(x)), f, df)
    }
    x
}
```

Output:
```
defined [32mfunction[39m [36msolveEquationNewtonRaphson[39m
```

<!-- grade_id: t5c, points: 5, grade: true, solution: false -->
```scala
// Test Cases
def f1(x: Double) = 2.0 * math.pow(x,3) - math.pow(x, 2) - 1.0
def df1(x: Double) = 6.0 * math.pow(x, 2) - 2 * x 
val x1 = solveEquationNewtonRaphson(3.0, f1, df1)
assert(math.abs(f1(x1)) <= 1e-06, "1")

def f2(x: Double) = math.sin(x) - 0.5 * x 
def df2(x: Double) = math.cos(x) - 0.5
val x2 = solveEquationNewtonRaphson(0.7, f2, df2)
assert(math.abs(f2(x2)) <= 1e-06, "2")


def f3(x: Double) = math.log(x) - 2 
def df3(x: Double) = 1/x 
val x3 = solveEquationNewtonRaphson(2.0, f3, df3)
assert(math.abs(f3(x3)) <= 1e-06, "3")

def f4(x: Double) = math.atan(x) - 1
def df4(x: Double) = 1.0/(1.0 + x*x)
val x4 = solveEquationNewtonRaphson(3.0, f4, df4)
assert(math.abs(f4(x4)) <= 1e-06, "4")

def f5(x: Double) = math.cos(x)* math.cos(x) - 1
def df5(x: Double) = -2.0 * math.sin(x) * math.cos(x)
val x5 = solveEquationNewtonRaphson(3.5, f5, df5)
assert(math.abs(f5(x5)) <= 1e-06, "5")
```

Output:
```
defined [32mfunction[39m [36mf1[39m
defined [32mfunction[39m [36mdf1[39m
[36mx1[39m: [32mDouble[39m = [32m1.000000000001152[39m
defined [32mfunction[39m [36mf2[39m
defined [32mfunction[39m [36mdf2[39m
[36mx2[39m: [32mDouble[39m = [32m8.537195744177245E-13[39m
defined [32mfunction[39m [36mf3[39m
defined [32mfunction[39m [36mdf3[39m
[36mx3[39m: [32mDouble[39m = [32m7.389056098804318[39m
defined [32mfunction[39m [36mf4[39m
defined [32mfunction[39m [36mdf4[39m
[36mx4[39m: [32mDouble[39m = [32m1.5574077246540006[39m
defined [32mfunction[39m [36mf5[39m
defined [32mfunction[39m [36mdf5[39m
[36mx5[39m: [32mDouble[39m = [32m3.1416751165142625[39m
```

```scala

```