# CSCI 3155: Assignment 2 With Solutions

Topics covered: recursion and inductive definitions.

__Name__: Julia Sanford

## Problem 1 (10 points total)

### A (5 points) Recursive `logTwo`
Given a positive integer $n$, we wish to compute $\log_2(n)$. Write a recursive function `logTwo(n: Int): Int` that does the same. Ensure that your precondition restricts the inputs to  $n > 0$.

<!-- grade_id: answer_1a, points: 0, grade: false, solution: true -->
```scala
def logTwo(n: Int): Int = { 
    // YOUR CODE HERE
    require(n > 0)
    if (n < 2){
        return 0
    }
    return 1 + logTwo(n / 2)
}
```

Output:
```
defined [32mfunction[39m [36mlogTwo[39m
```

<!-- grade_id: tests_1a, points: 5, grade: true, solution: false -->
```scala
// BEGIN TESTS
assert(logTwo(1) == 0, "logTwo(1) must be 0")
assert(logTwo(2) == 1, "logTwo(2) must be 1")
assert(logTwo(3) == 1, "logTwo(3) must be 1")
assert(logTwo(15) == 3, "logTwo(15) must be 3")
assert(logTwo(129) == 7, "logTwo(129) must be 8")
print("Tests Passed")
// END TESTS
```

Output:
```
Tests Passed
```

### B (5 points) _Tail_ Recursive `logTwo`
Now, write a tail recursive version of the `logTwo` function, called `logTwoTail`. You should implement the helper function as a tail recursive function.

<!-- grade_id: answer_1b, points: 0, grade: false, solution: true -->
```scala
import scala.annotation.tailrec

@tailrec
def logTwoHelper(n: Int, acc: Int): Int = {
    // YOUR CODE HERE
    require(n > 0)
    if (n < 2){
        return acc
    }
    return logTwoHelper(n / 2, 1 + acc)
}

def logTwoTail(n: Int): Int = { 
    // YOUR CODE HERE
    return logTwoHelper(n, 0)
}

```

Output:
```
[32mimport [39m[36mscala.annotation.tailrec

[39m
defined [32mfunction[39m [36mlogTwoHelper[39m
defined [32mfunction[39m [36mlogTwoTail[39m
```

<!-- grade_id: tests_1b, points: 5, grade: true, solution: false -->
```scala
// BEGIN TEST
assert(logTwoTail(1) == 0, "logTwo(1) must be 0")
assert(logTwoTail(2) == 1, "logTwo(2) must be 1")
assert(logTwoTail(3) == 1, "logTwo(3) must be 1")
assert(logTwoTail(15) == 3, "logTwo(15) must be 3")
assert(logTwoTail(129) == 7, "logTwo(129) must be 8")
print("Tests Passed")
//END TEST
```

Output:
```
Tests Passed
```

## Problem 2 (15 points)

We define a function `shuffleString(s: String): String` that, given an
input string `s`, does the following:
- If the string `s` is empty or length 1, the result is the same as input string `s`.
- Let n be the length of s.
- Recursively call `shuffleString` on the substring `s(n/2)... s(n-1)`. Let `s1` be the result
- Concatenate `s1` to the first half `s(0)..s(n/2-1)` to the result of the call.

Here is an implementation of this function and some examples, for your reference.


```scala
def shuffleString(s: String): String = {
    val n = s.length()
    if (n <= 1) { s }
    else {
        val secondHalf = s.substring(n/2, n)
        val shuffledHalf = shuffleString(secondHalf)
        val firstHalf = s.substring(0, n/2)
        return  shuffledHalf +  firstHalf 
    }
}

val f0 = shuffleString("1234")
val f1 = shuffleString("12345")
val f2 = shuffleString("123456")
val f3 = shuffleString("1234567")
val f4 = shuffleString("12345678")
val f5 = shuffleString("123456789")
```

Output:
```
defined [32mfunction[39m [36mshuffleString[39m
[36mf0[39m: [32mString[39m = [32m"4312"[39m
[36mf1[39m: [32mString[39m = [32m"54312"[39m
[36mf2[39m: [32mString[39m = [32m"654123"[39m
[36mf3[39m: [32mString[39m = [32m"7645123"[39m
[36mf4[39m: [32mString[39m = [32m"87561234"[39m
[36mf5[39m: [32mString[39m = [32m"987561234"[39m
```

Implement a tail recursive version of `shuffleString` using an accumulator variable. It will help to carefully examine how different parts of the string get rearranged to design this function or even first try to write a simple while loop that mimics `shuffleString`.

<!-- grade_id: answer_2, points: 0, grade: false, solution: true -->
```scala
def tailRecursiveShuffleString(s: String, acc: String = ""): String = {
    // YOUR CODE HERE
    val n = s.length()
    if (n <= 1) { s + acc }
    else {
        val firstHalf = s.substring(0, n/2)
        val secondHalf = s.substring(n/2, n)
        return tailRecursiveShuffleString(secondHalf, firstHalf + acc)
    }
}
```

Output:
```
defined [32mfunction[39m [36mtailRecursiveShuffleString[39m
```

<!-- grade_id: tests_2, points: 15, grade: true, solution: false -->
```scala
// BEGIN TEST
assert(shuffleString("1") == tailRecursiveShuffleString("1"), "Failed test 1")
assert(shuffleString("12") == tailRecursiveShuffleString("12"), "Failed test 12")
assert(shuffleString("123") == tailRecursiveShuffleString("123"), "Failed test 123")
assert(shuffleString("1234") == tailRecursiveShuffleString("1234"), "Failed test 1234")
assert(shuffleString("12345") == tailRecursiveShuffleString("12345"), "Failed test 12345")
assert(shuffleString("123456") == tailRecursiveShuffleString("123456"), "Failed test 123456")
assert(shuffleString("1234567") == tailRecursiveShuffleString("1234567"), "Failed test 1234567")
assert(shuffleString("12345678") == tailRecursiveShuffleString("12345678"), "Failed test 12345678")
assert(shuffleString("123456789") == tailRecursiveShuffleString("123456789"), "Failed test 123456789")
assert(shuffleString("1234567890") == tailRecursiveShuffleString("1234567890"), "Failed test 1234567890")
print("Tests Passed")
// END TEST
```

Output:
```
Tests Passed
```

## Problem 3 (20 points)

Convert the following inductive definition for regular expressions into a grammar first and then into a set of scala classes.

A regular expression is defined inductively as follows:
- Any string s is an "atomic" regular expression.
- If $r_1$, $r_2$ are regular expressions then so are
  - The concatenation $r_1 ; r_2$, 
  - The disjunction $r_1 | r_2$, and 
  - The conjunction $r_1 \& r_2$.
- If $r$ is a regular expression, then its Kleene star $r^*$ is also a regular expression.

Use the constructor symbols $\text{Atom}(s)$ to denote an atomic regular expression, $\text{Concat}(r_1, r_2)$ for the ";" operator, $\text{Or}(r_1, r_2)$ for the "|" operator, $\text{And}(r_1, r_2)$ for the "&" operator and $\text{Star}(r)$ for the Kleene-star operator. You may use
the nonterminal $\textbf{string}$ without definition to denote a string of characters.




### A (7 point)
Write the grammar using constructor symbols for the inductive definition above. __Tip:__ you can examine the notebooks with inductive definitions to see how we typeset the grammar. There are no tests for this because it will be manualy graded.

<!-- grade_id: answer_3a, points: 7, grade: true, solution: true -->
$$
\begin{array}{rcl}
    \textbf{Regex} & \rightarrow & Atom(string) \\
                 &           | & Concat(\textbf{Regex}, \textbf{Regex}) \\
                 &           | & Or(\textbf{Regex}, \textbf{Regex}) \\
                 &           | & And(\textbf{Regex}, \textbf{Regex}) \\
                 &           | & Star(\textbf{Regex}) \\
\end{array}
$$

### B (7 points)
Define the structure as a set of case class in scala.

__NOTICE__: Tests are hidden.

<!-- grade_id: answer_3b, points: 0, grade: false, solution: true -->
```scala
sealed trait Regex
// Use constructors: Atom, Concat, Or, And and Star
// YOUR CODE HERE
case class Atom(s: String) extends Regex
case class Concat(r1: Regex, r2: Regex) extends Regex
case class Or(r1: Regex, r2: Regex) extends Regex
case class And(r1: Regex, r2: Regex) extends Regex
case class Star(r1: Regex) extends Regex
```

Output:
```
defined [32mtrait[39m [36mRegex[39m
defined [32mclass[39m [36mAtom[39m
defined [32mclass[39m [36mConcat[39m
defined [32mclass[39m [36mOr[39m
defined [32mclass[39m [36mAnd[39m
defined [32mclass[39m [36mStar[39m
```

<!-- grade_id: tests_3b, points: 7, grade: true, solution: false -->
```scala
// WARNING: Tests are hidden because they directly give away the
//   answer.
// The tests will ensure there is no compile error when we try
//   to use the classes you were asked to define.
```

### C (6 points)
Write down the representation of the regular expression in Scala. Your cell must define a term that should be called `finalAnswerC`.

$$ (\text{"hello"})^* ;  ((\text{"scala"}; \text{"best"})^*) $$

__NOTICE__: Tests are hidden.

<!-- grade_id: answer_3c, points: 0, grade: false, solution: true -->
```scala
val finalAnswerC = {
    // YOUR CODE HERE
    val s1 = Star(Atom("hello"))
    val s2 = Star(Concat(Atom("scala"), Atom("best")))
    Concat(s1, s2)
}
```

Output:
```
[36mfinalAnswerC[39m: [32mConcat[39m = Concat(Star(Atom(hello)),Star(Concat(Atom(scala),Atom(best))))
```

<!-- grade_id: tests_3c, points: 6, grade: true, solution: false -->
```scala
// WARNING: Tests are hidden because they directly give away the
//   answer.
// The tests will do an equality check with the expected solution.
```