# CSCI 3155: Assignment 3 

Topics covered: ASTs, pattern matching, higher order functions.

__Name__: Julia Sanford

## Problem 1 (15 points)

Develop a switch statement for the `mython` (var) language developed in class. 
The syntax for the to be added switch statement is given in terms of additional rules that
add to the existing definitions in the class notes.

$$\begin{array}{rcl}
\mathbf{Statement} & \rightarrow & \text{Switch}(\mathbf{Expr},  \mathbf{Case}* ) \\[5pt]
\mathbf{Case} & \rightarrow & \text{SwitchCase}(\mathbf{Expr}, \mathbf{Statement}* ) \\[5pt]
& | & \text{DefaultCase}(\mathbf{Statement}*) \\
\end{array}$$

The existing scala language definitions for `mython` are recalled for below from the class notes.


Augment the definitions by adding the extra definitions to support switch statements. Please use
the same constructor names as in the grammar: `Switch` for the switch statement and 
`SwitchCase` and `DefaultCase` for the two rules corresponding to the __Case__ specifications.

For your convenience, the definitions used in the class are given. Add your solution at the end.


<!-- grade_id: cell-a3a78387d2a93841, points: 0, grade: false, solution: true -->
```scala
sealed trait Expr
case class Const(d: Double) extends Expr 
// 1. We cheated a bit and allowed all floating point numbers
// Also, this deviates from the grammar
case class Ident(s: String) extends Expr
// 2. We allow any string to be an identifier for now instead of the regular expression shown in the grammar.
case class Plus(e1: Expr, e2: Expr ) extends Expr
case class Minus(e1: Expr, e2: Expr) extends Expr
case class Mult(e1: Expr, e2: Expr) extends Expr
case class Div(e1: Expr, e2: Expr) extends Expr
case class Negate(e: Expr) extends Expr
case class Log(e: Expr) extends Expr
case class Exp(e: Expr) extends Expr
case class Sine(e: Expr) extends Expr
case class Cosine(e: Expr) extends Expr

sealed trait CondExpr
case object ConstTrue extends CondExpr
case object ConstFalse extends CondExpr
case class Geq(e1: Expr, e2: Expr) extends CondExpr
case class Leq(e1: Expr, e2: Expr) extends CondExpr
case class Eq(e1: Expr, e2: Expr) extends CondExpr
case class And(c1: CondExpr, c2: CondExpr) extends CondExpr
case class Or(c1: CondExpr, c2: CondExpr) extends CondExpr
case class Not(c: CondExpr) extends CondExpr

sealed trait Declaration
sealed trait Statement
sealed trait MythonProgram

case class Program(decls: List[Declaration], stmts: List[Statement], returnAtEnd: Expr) extends MythonProgram // We stripped the ReturnStmt tag since it is redundant
case class VarDecl(identifier: String, rhsExpr: Expr) extends Declaration
case class AssignStmt(identifier: String, rhsExpr: Expr) extends Statement
case class WhileStmt(cond: CondExpr, stmts: List[Statement]) extends Statement
case class IfThenElseStmt(cond: CondExpr, stmtsThen: List[Statement], stmtsElse: List[Statement]) extends Statement
case class ReturnStmt(retExpr: Expr) extends Statement

// YOUR CODE HERE
sealed trait Case
case class Switch(e: Expr, cases: List[Case]) extends Statement
case class SwitchCase(e: Expr, stmts: List[Statement]) extends Case
case class DefaultCase(stmts: List[Statement]) extends Case
```

Output:
```
defined [32mtrait[39m [36mExpr[39m
defined [32mclass[39m [36mConst[39m
defined [32mclass[39m [36mIdent[39m
defined [32mclass[39m [36mPlus[39m
defined [32mclass[39m [36mMinus[39m
defined [32mclass[39m [36mMult[39m
defined [32mclass[39m [36mDiv[39m
defined [32mclass[39m [36mNegate[39m
defined [32mclass[39m [36mLog[39m
defined [32mclass[39m [36mExp[39m
defined [32mclass[39m [36mSine[39m
defined [32mclass[39m [36mCosine[39m
defined [32mtrait[39m [36mCondExpr[39m
defined [32mobject[39m [36mConstTrue[39m
defined [32mobject[39m [36mConstFalse[39m
defined [32mclass[39m [36mGeq[39m
defined [32mclass[39m [36mLeq[39m
defined [32mclass[39m [36mEq[39m
defined [32mclass[39m [36mAnd[39m
defined [32mclass[39m [36mOr[39m
defined [32mclass[39m [36mNot[39m
defined [32mtrait[39m [36mDeclaration[39m
defined [32mtrait[39m [36mStatement[39m
defined [32mtrait[39m [36mMythonProgram[39m
defined [32mclass[39m [36mP
```

<!-- grade_id: cell-b4433298fac96b06, points: 5, grade: true, solution: false -->
```scala
// WARNING: Tests are hidden because they directly give away the
//   answer. Feel free to come into office hours to ask whether
//   you're on the right track.
// The tests will create an AST with a SwitchCase.
```

<!-- grade_id: cell-590403a05cbd399c, points: 5, grade: true, solution: false -->
```scala
// More hidden tests for DefaultCase
```

<!-- grade_id: cell-73f9de24de44bb2f, points: 5, grade: true, solution: false -->
```scala
// More hidden tests for Switch
```

## Problem 2 (15 Points)

Add list manipulation expressions to the allowed expressions in the `mython` language. 

We allow list expressions (use a new non terminal called __ListExpr__) that can do the following:
- the empty list: use constructor symbol `EmptyList`
- explicitly create a list of arithmetic expressions: use constructor symbol `NewList`
- Cons an arithmetic expression to the front of a list: use constructor symbol `Cons`
- Concatenate two lists defined by list expressions: use constructor symbol `Concat`

Next, we add to conditional expressions the following:
- A condition expression to check if a list is empty:  Use constructor `IsEmptyList`

Finally, we add a special foreach loop statement over the elements of a list.

`foreach (x in list_expr) begin  .. list of statements...  end `

- Add a statement to loop over the elements of a list: use constructor`ForEach`

Write down the extra grammar rules that you would add to the overall 
`mython` grammar to create these functionalities.

<!-- grade_id: cell-76cfd1d2c6a5898d, points: 15, grade: true, solution: true -->
$$
\begin{array}{rcl}
    \textbf{ListExpr} & \rightarrow & EmptyList \\
                      &           | & NewList(\textbf{Expr*}) \\
                      &           | & Cons(\textbf{Expr}, \textbf{ListExpr}) \\
                      &           | & Concat(\textbf{ListExpr}, \textbf{ListExpr}) \\
    \textbf{CondExp}  & \rightarrow & IsEmptyList(\textbf{ListExpr}) \\
                      &           | & ForEach(\textbf{ListExpr}, \textbf{Statement*}) \\
\end{array}
$$

## Problem 3 (20 points)

__(A, 5 points)__ We defined lists in the class. Write a recursive procedure to get the nth element of the list or throw an `IllegalArgumentException` if $n < 0$ or $n >= length of list $. Assume $n=0$ obtains the very first element
and $n = length of list -1$ yields very last element. 





<!-- grade_id: cell-724b4ca89445ceaf, points: 0, grade: false, solution: true -->
```scala
sealed trait NumList
case object Nil extends NumList
case class Cons(n: Int, l: NumList) extends NumList 

def listLength(lst: NumList): Int = { 
    lst match { // pattern match the lst
        case Nil => 0
        case Cons(_, tl) => 1 + listLength(tl) // _ here is to tell scala that I do not care what is the element.
    }
}

def getNthElement(lst: NumList, n: Int): Int = {
    // YOUR CODE HERE
    val l = listLength(lst)
    lst match {
        case _ if(n < 0 || n >= l) => throw new IllegalArgumentException("n is negative")
        case Cons(_, rest) if (n != 0) => getNthElement(rest, n-1)
        case Cons(ret, rest) if (n == 0) => ret
    }
}
```

Output:
```
defined [32mtrait[39m [36mNumList[39m
defined [32mobject[39m [36mNil[39m
defined [32mclass[39m [36mCons[39m
defined [32mfunction[39m [36mlistLength[39m
defined [32mfunction[39m [36mgetNthElement[39m
```

<!-- grade_id: cell-2bc6d5ee94793b48, points: 5, grade: true, solution: false -->
```scala
// BEGIN TESTS
val l1 = Nil
val l2 = Cons(1, Cons(-1, Nil))
val l3 = Cons(1, Cons(2, l2))
val l4 = Cons(0, Cons(4, Cons(8, l3)))

val test1 = try {
    getNthElement(Nil, 3);
    assert(false, "Test 1 : getNthElement(Nil, 3) should raise an IllegalArgumentException. Your code did not.")
} catch {
    case e: IllegalArgumentException => "OK"
} 

assert(getNthElement(l2, 0) == 1, "Test2: getNthElement(l2, 0)  failed (expected answer = 1)")
assert(getNthElement(l3, 3) == -1, "Test3: getNthElement(l3, 3)  failed (expected answer = -1)")
assert(getNthElement(l4, 2) == 8, "Test4: getNthElement(l4, 2)  failed (expected answer = 8)")

val test2 = try {
    getNthElement(l4, 8);
    assert(false, "Test 5 : getNthElement(l4, 8) should raise an IllegalArgumentException. Your code did not.")
} catch {
    case e: IllegalArgumentException => "OK"
} 

// END TESTS
```

Output:
```
[36ml1[39m: [32mNil[39m.type = Nil
[36ml2[39m: [32mCons[39m = Cons(1,Cons(-1,Nil))
[36ml3[39m: [32mCons[39m = Cons(1,Cons(2,Cons(1,Cons(-1,Nil))))
[36ml4[39m: [32mCons[39m = Cons(0,Cons(4,Cons(8,Cons(1,Cons(2,Cons(1,Cons(-1,Nil)))))))
[36mtest1[39m: [32mAny[39m = OK
[36mtest2[39m: [32mAny[39m = OK
```

__(B, 7 points)__ Write a recursive procedure that returns true if the list has the Fibonacci property. I.e, every element
at position $i \geq 2$ is the sum of the two preceding elements. Note that the property is trivially true for lists of sizes $0$ and $1$.

<!-- grade_id: cell-eedd55e1f86bd6c6, points: 0, grade: false, solution: true -->
```scala
def isFibonacciList(lst: NumList): Boolean = {
    // YOUR CODE HERE
    val l = listLength(lst)
    lst match {
        case _ if(l <= 2) => true
        case Cons(one, Cons(two, Cons(three, rest))) if(one + two == three) => {
            isFibonacciList(Cons(two, Cons(three, rest)))
        }
        case Cons(one, Cons(two, Cons(three, rest))) if(one + two != three) => {false}
    }
}
```

Output:
```
defined [32mfunction[39m [36misFibonacciList[39m
```

<!-- grade_id: cell-33a4f21c01f0c7c4, points: 7, grade: true, solution: false -->
```scala
// BEGIN TESTS
val l1 = Cons(12, Cons(25, Cons(37, Nil)))
assert(isFibonacciList(l1), 
       "Test case 1 :  isFibonacciList(l1) -- should return true")

val l2 = Cons(14, Cons(-1, Cons(13, l1 )))
assert(isFibonacciList(l2), 
       "Test case 2 :  isFibonacciList(l2) -- should return true")

val l3 = Cons(7, Cons(7, l2))

assert(!isFibonacciList(l3), 
       "Test case 3 :  isFibonacciList(l3) -- should return false")

val l4 = Cons(0, Cons(0, Cons(0, Cons(0, Cons(0, Cons(0, Nil))))))
assert(isFibonacciList(l2), 
       "Test case 4:  isFibonacciList(l4) -- should return true")

// END TESTS
```

Output:
```
[36ml1[39m: [32mCons[39m = Cons(12,Cons(25,Cons(37,Nil)))
[36ml2[39m: [32mCons[39m = Cons(14,Cons(-1,Cons(13,Cons(12,Cons(25,Cons(37,Nil))))))
[36ml3[39m: [32mCons[39m = Cons(7,Cons(7,Cons(14,Cons(-1,Cons(13,Cons(12,Cons(25,Cons(37,Nil))))))))
[36ml4[39m: [32mCons[39m = Cons(0,Cons(0,Cons(0,Cons(0,Cons(0,Cons(0,Nil))))))
```

__(C, 8 points)__ Write a recursive function `filterNumList(l: NumList, f: Int => Boolean): NumList` that takes in a `NumList` and a function `f: Int => Boolean`. 

1. It should return a new list that consist of all elements of the list `l` that return `true` when the function `f` is called on them.
2. The returned list elements must preserve the same order as in the original list.




<!-- grade_id: cell-0f9f5a51ae34b3bb, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
def filterNumList(l: NumList, f: Int => Boolean): NumList = {
    l match {
        case Nil => Nil
        case Cons(first, rest) => {
            if (f(first)){
                Cons(first, filterNumList(rest, f))
            }
            else {
                filterNumList(rest, f)
            }
        }
        
    }
}
```

Output:
```
defined [32mfunction[39m [36mfilterNumList[39m
```

<!-- grade_id: cell-a2549ef77f3f9bb0, points: 8, grade: true, solution: false -->
```scala
// BEGIN TESTS
val l1 = Cons(12, Cons(25, Cons(37, Nil)))
def f1(j: Int): Boolean =  j <= 25 && j >= 12
assert(filterNumList(l1, f1) == Cons(12, Cons(25, Nil)), "Test 1 failed.")


val l2 = Cons(22, Cons(135, Cons(137, l1)))
def f2(j: Int): Boolean =  j % 5 == 0
assert(filterNumList(l2, f2) == Cons(135, Cons(25, Nil)), "Test 2 failed.")

def f3(j: Int): Boolean =  j >= 210
assert(filterNumList(l2, f3) == Nil, "Test 3 failed.")
assert(filterNumList(Nil, f3) == Nil, "Test 4 failed.")

val l4 = Cons(0, Cons(0, Cons(0, Cons(0, Cons(0, Cons(0, Nil))))))

def f4(j: Int): Boolean =  j <= 0
assert(filterNumList(l4, f4) == l4, "Test 5 failed")

// END TESTS
```

Output:
```
[36ml1[39m: [32mCons[39m = Cons(12,Cons(25,Cons(37,Nil)))
defined [32mfunction[39m [36mf1[39m
[36ml2[39m: [32mCons[39m = Cons(22,Cons(135,Cons(137,Cons(12,Cons(25,Cons(37,Nil))))))
defined [32mfunction[39m [36mf2[39m
defined [32mfunction[39m [36mf3[39m
[36ml4[39m: [32mCons[39m = Cons(0,Cons(0,Cons(0,Cons(0,Cons(0,Cons(0,Nil))))))
defined [32mfunction[39m [36mf4[39m
```